**Important notes:**

**Important notes:**

- You can complete this exercise in 1) Colab or 2) on your local machine
  1. **Colab**: Click on the 🚀 symbol at the top of the page and select Colab. When you finished the exercise, download the file: `File` > `Download` > `ipynb`. 
  2. **Local**: click on the download button at the top of the page and choose `.ipynb`. Activate the conda environment `mr` before you start: `conda activate mr`


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.


- The function **NotImplementedError()** prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - in *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - if you use *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 


Good luck!

```python
NAME = "ns166"
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Find the best model

## Setup

```python
import pandas as pd
import altair as alt
import warnings

warnings.simplefilter(action='ignore', category=FutureWarning)
alt.data_transformers.disable_max_rows()
```

Output:
```
DataTransformerRegistry.enable('default')
```

## Data

```python
# let's change some values in ads
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

## Analysis

```python
alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales',
    tooltip= ['ads', 'sales']
).interactive()
```

Output:
```
alt.Chart(...)
```

*Is there a "perfect" relationship between the two variables (like in our last example)?*

## Model

Remember our simple model from before. Instead of using `number_0` and `number_1`, we use the following expressions:

- `number_0` = `intercept`

- `number_1` = `slope` (also called coefficient)


Use the following values to make a prediction:

- intercept = 500
- slope = 2

Hint:

---


```python
intercept = ___ 
slope = ___

df['___'] = ___ + ___ * df['___'] 
```

---

- Make a prediction for `sales` by using the variable `ads`
- Save the new column as `sales_prediction`

<!-- grade_id: sales-prediction, points: 0, grade: false, solution: true -->
```python
intercept = 500 
slope = 2

df['sales_prediction'] = intercept + slope * df['ads']
```

```python
df.head()
```

Output:
```
sales   ads  sales_prediction
0   2500   900              2300
1   4500  1400              3300
2   6500  3600              7700
3   8500  3800              8100
4  10500  6200             12900
```

<!-- grade_id: c-sales-prediction, points: 4, grade: true, solution: false -->
```python
# Check your code
assert df.loc[0, 'sales_prediction'] == 2300 
```

```python
chart = alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales'
)

line = alt.Chart(df).mark_line().encode(
         alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
         alt.Y('sales_prediction', axis=alt.Axis(title="Sales (in units)")),
         color=alt.value("#0001F5"))

chart + line
```

Output:
```
alt.LayerChart(...)
```

*Can you guess why the two numbers are called intercept and slope?* Write down your answer below.

Intercept ist der Schnittpunkt mit der y-Achse bei x=0.
Slope ist die Steigung des y-Wertes auf der x-Achse.