**Important notes:**

**Important notes:**

- You can complete this exercise in 1) Colab or 2) on your local machine
  1. **Colab**: Click on the 🚀 symbol at the top of the page and select Colab. When you finished the exercise, download the file: `File` > `Download` > `ipynb`. 
  2. **Local**: click on the download button at the top of the page and choose `.ipynb`. Activate the conda environment `mr` before you start: `conda activate mr`


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.


- The function **NotImplementedError()** prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - in *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - if you use *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 


Good luck!

```python
NAME = "jk251"
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Find the best model

## Setup

```python
import pandas as pd
import altair as alt
import warnings

warnings.simplefilter(action='ignore', category=FutureWarning)
alt.data_transformers.disable_max_rows()
```

Output:
```
DataTransformerRegistry.enable('default')
```

## Data

```python
# let's change some values in ads
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

## Analysis

```python
alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales',
    tooltip= ['ads', 'sales']
).interactive()
```

Output:
```
alt.Chart(...)
```

*Is there a "perfect" relationship between the two variables (like in our last example)?*

## Model

Remember our simple model from before. Instead of using `number_0` and `number_1`, we use the following expressions:

- `number_0` = `intercept`

- `number_1` = `slope` (also called coefficient)


Use the following values to make a prediction:

- intercept = 500
- slope = 2

Hint:

---


```python
intercept = ___ 
slope = ___

df['___'] = ___ + ___ * df['___'] 
```

---

- Make a prediction for `sales` by using the variable `ads`
- Save the new column as `sales_prediction`

<!-- grade_id: sales-prediction, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df['sales']= 500 + 2 * df['ads'] 
sales_prediction = df 
```

```python
df.head()
```

Output:
```
sales   ads
0   2300   900
1   3300  1400
2   7700  3600
3   8100  3800
4  12900  6200
```

<!-- grade_id: c-sales-prediction, points: 4, grade: true, solution: false -->
```python
# Check your code
assert df.loc[0, 'sales_prediction'] == 2300 
```

Output:
```
File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/frame.py:4272, in DataFrame._get_item_cache(self, item)
   4267 res = cache.get(item)
   4268 if res is None:
   4269     # All places that call _get_item_cache have unique columns,
   4270     #  pending resolution of GH#33047
-> 4272     loc = self.columns.get_loc(item)
   4273     res = self._ixs(loc, axis=1)
   4275     cache[item] = res

File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/indexes/base.py:3802, in Index.get_loc(self, key, method, tolerance)
   3800     return self._engine.get_loc(casted_key)
   3801 except KeyError as err:
-> 3802     raise KeyError(key) from err
   3803 except TypeError:
   3804     # If we have a listlike key, _check_indexing_error will raise
   3805     #  InvalidIndexError. Otherwise we fall through and re-raise
   3806     #  the TypeError.
   3807     self._check_indexing_error(key)

KeyError: 'sales_prediction'
```

```python
chart = alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales'
)

line = alt.Chart(df).mark_line().encode(
         alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
         alt.Y('sales_prediction', axis=alt.Axis(title="Sales (in units)")),
         color=alt.value("#0001F5"))

chart + line
```

Output:
```
File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/altair/utils/schemapi.py:56, in _todict(obj, validate, context)
     54 """Convert an object to a dict representation."""
     55 if isinstance(obj, SchemaBase):
---> 56     return obj.to_dict(validate=validate, context=context)
     57 elif isinstance(obj, (list, tuple, np.ndarray)):
     58     return [_todict(v, validate, context) for v in obj]

File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/altair/vegalite/v4/schema/channels.py:40, in FieldChannelMixin.to_dict(self, validate, ignore, context)
     38 elif not (type_in_shorthand or type_defined_explicitly):
     39     if isinstance(context.get('data', None), pd.DataFrame):
---> 40         raise ValueError("{} encoding field is specified without a type; "
     41                          "the type cannot be inferred because it does not "
     42                          "match any column in the data.".format(shorthand))
     43     else:
     44         raise ValueEr
```

*Can you guess why the two numbers are called intercept and slope?* Write down your answer below.