```python
NAME = "jk251"
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Find the best model

## Setup

```python
import pandas as pd

from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
from sklearn.metrics import mean_absolute_error
```

## Data

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

### Variable lists

We need to prepare our data for the scikit-learn model (which does not work with Pandas dataframes):

- We create the variable `y_label` (our dependent variable, also called label): choose "`sales`" (this step will be helpful in later notebooks)

- We select our features (independent variables) and save them as `X`: choose "`ads`"

- Select and save the dependent variable as `y`

Hint:

```pyhton
y_label = "___"

X = df[["___"]]
y = df[___]
```

<!-- grade_id: y_label, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y_label = "sales"

x = df[["ads"]]
y = df[y]
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [5], line 5
      2 y_label = "sales"
      4 x = df[["ads"]]
----> 5 y = df[y]

NameError: name 'y' is not defined
```

```python
X
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [6], line 1
----> 1 X

NameError: name 'X' is not defined
```

```python
y
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [7], line 1
----> 1 y

NameError: name 'y' is not defined
```

<!-- grade_id: c, points: 2, grade: true, solution: false -->
```python
### Check your code (note that there are some hidden checks)
assert len(X) == 10

### BEGIN HIDDEN TESTS 
assert y_label == "sales"
assert len(y) == 10
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [8], line 2
      1 ### Check your code (note that there are some hidden checks)
----> 2 assert len(X) == 10
      4 ### BEGIN HIDDEN TESTS 
      5 assert y_label == "sales"

NameError: name 'X' is not defined
```

## Model

### Select model

Hint:

Save the `LinearRegression()`estimator as object called `reg`

---

```pyhton
___ = ___()
```

---

<!-- grade_id: reg, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [9], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: c-reg, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.get_params() == {'copy_X': True, 'fit_intercept': True, 'n_jobs': None, 'normalize': 'deprecated', 'positive': False}
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [10], line 2
      1 ### Check your code
----> 2 assert reg.get_params() == {'copy_X': True, 'fit_intercept': True, 'n_jobs': None, 'normalize': 'deprecated', 'positive': False}

NameError: name 'reg' is not defined
```

### Fit model


Fit the model (your `reg` object) to the data with `.fit()`.

In the fit function, you first need to provide the features and then the dependent variable (features, dependent variable)

Hint:


---

```pyhton
___.___(___, ___)
```

---



<!-- grade_id: reg-fit, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [11], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: c-reg-fit, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.positive == False
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [12], line 2
      1 ### Check your code
----> 2 assert reg.positive == False

NameError: name 'reg' is not defined
```

### Coefficients

```python
# Intercept
reg.intercept_
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [13], line 2
      1 # Intercept
----> 2 reg.intercept_

NameError: name 'reg' is not defined
```

```python
# Slope
reg.coef_
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [14], line 2
      1 # Slope
----> 2 reg.coef_

NameError: name 'reg' is not defined
```

### Make predictions


Use your model to:

- predict sales from your ads (`X`) values with the method `.predict()`

- Save the result as `y_pred`

Hint:

---

```pyhton
___ = ___.___(___)
```

---



<!-- grade_id: y-pred, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [15], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
y_pred[0]
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [16], line 1
----> 1 y_pred[0]

NameError: name 'y_pred' is not defined
```

<!-- grade_id: c-y-pred, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 2836 < y_pred[0] <  2840
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [17], line 2
      1 ### Check your code
----> 2 assert 2836 < y_pred[0] <  2840

NameError: name 'y_pred' is not defined
```

### Evaluation

#### Mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the mean squared error and save the result as `mse`

Hint:

---

```pyhton
___ = ___(___, ___)
```

---

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [18], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
print(f'The mean squared error equals {mse:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [19], line 1
----> 1 print(f'The mean squared error equals {mse:.0f} ')

NameError: name 'mse' is not defined
```

<!-- grade_id: c-mse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1160738 < mse < 1160740 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [20], line 2
      1 ### Check your code
----> 2 assert 1160738 < mse < 1160740

NameError: name 'mse' is not defined
```

#### Root mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the root mean squared error and save the result as `rmse`.

Hint:

---

```pyhton
___ = ___(___, ___, ___)
```

- use the option `squared=False` 

---

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [21], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
print(f'The root mean squared error equals {rmse:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [22], line 1
----> 1 print(f'The root mean squared error equals {rmse:.0f} ')

NameError: name 'rmse' is not defined
```

<!-- grade_id: c-rmse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1076 < rmse < 1078 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [23], line 2
      1 ### Check your code
----> 2 assert 1076 < rmse < 1078

NameError: name 'rmse' is not defined
```

#### Mean absolute error

Use the scikit-learn function `mean_absolute_error()` to calculate the mean absolute error and save the result as `mae`.

Hint:

---

```pyhton
___ = ___(___, ___)
```


---

<!-- grade_id: mae, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [24], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
print(f'The mean absolute error equals {mae:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [25], line 1
----> 1 print(f'The mean absolute error equals {mae:.0f} ')

NameError: name 'mae' is not defined
```

<!-- grade_id: c-mae, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 885 < mae < 887 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [26], line 2
      1 ### Check your code
----> 2 assert 885 < mae < 887

NameError: name 'mae' is not defined
```