**Important notes:**

**Important notes:**

- You can complete this exercise in 1) Colab or 2) on your local machine
  1. **Colab**: Click on the 🚀 symbol at the top of the page and select Colab. When you finished the exercise, download the file: `File` > `Download` > `ipynb`. 
  2. **Local**: click on the download button at the top of the page and choose `.ipynb`. Activate the conda environment `mr` before you start: `conda activate mr`


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.


- The function **NotImplementedError()** prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - in *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - if you use *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 


Good luck!

```python
NAME = ""
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Mean squared error

## Setup

We need the following modules:

- pandas
- scikit-learn

```python
import pandas as pd
from sklearn.metrics import mean_squared_error
```

## Data

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

### Data structure

```python
df
```

Output:
```
sales    ads
0   2500    900
1   4500   1400
2   6500   3600
3   8500   3800
4  10500   6200
5  12500   5200
6  14500   6800
7  16500   8300
8  18500   9800
9  20500  10100
```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 10 entries, 0 to 9
Data columns (total 2 columns):
 #   Column  Non-Null Count  Dtype
---  ------  --------------  -----
 0   sales   10 non-null     int64
 1   ads     10 non-null     int64
dtypes: int64(2)
memory usage: 288.0 bytes
```

## Model

```python
intercept = 500 
slope = 2

df['sales_prediction'] = intercept + slope * df['ads'] 
```

### Calculate mean squared error

---

Step 1: calculate the squared difference between true outcome and prediction:


The general formula:


$$Squared\ errors \ (SE) = (y_i-\hat{y_i})^2$$



For our data this means:

$$ se = (sales_i - sales\_prediction_i)^2$$

Hint

---

```python
df['___'] = (df['___'] - df['___']) ___ ___

```

---

- Use the `**` (power) operator to raise a value to the power of 2.

<!-- grade_id: se, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [8], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
df.head()
```

Output:
```
sales   ads  sales_prediction
0   2500   900              2300
1   4500  1400              3300
2   6500  3600              7700
3   8500  3800              8100
4  10500  6200             12900
```

<!-- grade_id: c-se, points: 2, grade: true, solution: false -->
```python
# Check your code
assert df.loc[0, 'se'] == 40000
```

Output:
```
File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/frame.py:4272, in DataFrame._get_item_cache(self, item)
   4267 res = cache.get(item)
   4268 if res is None:
   4269     # All places that call _get_item_cache have unique columns,
   4270     #  pending resolution of GH#33047
-> 4272     loc = self.columns.get_loc(item)
   4273     res = self._ixs(loc, axis=1)
   4275     cache[item] = res

File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/indexes/base.py:3802, in Index.get_loc(self, key, method, tolerance)
   3800     return self._engine.get_loc(casted_key)
   3801 except KeyError as err:
-> 3802     raise KeyError(key) from err
   3803 except TypeError:
   3804     # If we have a listlike key, _check_indexing_error will raise
   3805     #  InvalidIndexError. Otherwise we fall through and re-raise
   3806     #  the TypeError.
   3807     self._check_indexing_error(key)

KeyError: 'se'
```

---

Step 2: Sum up all squared errors

Calculate the sum of squared errors (SSE)

The general formula:

$$Sum\ of\ squared\ errors \ (SSE) = \sum_{i=1}^{n}(y_i-\hat{y_i})^2$$

which equals

$$Sum\ of\ squared\ errors \ (SSE) = \sum_{i=1}^{n} SE$$



For our data this means:



$$sse = \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2$$


which equals

$$sse = \sum_{i=1}^{n} se$$

Hint

Sum up the squared errors and save the result in a new object called `sse` (not in the Pandas dataframe)

---

```python
___ = df['___'].___
```

---

- Save the result as `sse`
- use the method `.sum()` to sum up all `se`

<!-- grade_id: sse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [11], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
sse
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [12], line 1
----> 1 sse

NameError: name 'sse' is not defined
```

<!-- grade_id: c-sse, points: 2, grade: true, solution: false -->
```python
# Chack your code
assert sse == 14520000
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [13], line 2
      1 # Chack your code
----> 2 assert sse == 14520000

NameError: name 'sse' is not defined
```

---

Step 3: Divide by the number of examples

The general formula:

$$Mean\ squared\ error \ (MSE) = \frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})^2$$

which equals

$$Mean\ squared\ error \ (MSE) = \frac{1}{n} \times SSE$$

*with n = the number of observations*



For our data this means:



$$mse = \frac{1}{n} \times \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2$$

which equals

$$mse = \frac{1}{n} \times sse $$


*with n = the number of observations*



```python
# Obtain number of examples (observations)
n = len(df)
```

```python
n
```

Output:
```
10
```

Calculate the mse and save the result as `mse`:

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [16], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
mse
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [17], line 1
----> 1 mse

NameError: name 'mse' is not defined
```

```python
# alternative way to print the result
print(f'The MSE equals {mse:.0f}')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [18], line 2
      1 # alternative way to print the result
----> 2 print(f'The MSE equals {mse:.0f}')

NameError: name 'mse' is not defined
```

<!-- grade_id: c-mse, points: 2, grade: true, solution: false -->
```python
# Check your code
assert 1451000 < mse < 1453000
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [19], line 2
      1 # Check your code
----> 2 assert 1451000 < mse < 1453000

NameError: name 'mse' is not defined
```

### Root mean squared error

---

Another important measure is the root mean squared error:

The general formula:

$$Root\ mean\ squared\ error \ (RMSE) = \sqrt{\frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})^2}$$

which equals

$$Root\ mean\ squared\ error \ (RMSE) = \sqrt{\frac{1}{n} \times SSE}$$

*with n = the number of observations*






For our data this means:


$$rmse = \sqrt{\frac{1}{n} \times \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2}$$

which equals

$$rmse = \sqrt{\frac{1}{n} \times sse} $$

which also equals

$$rmse = (\frac{1}{n} \times sse)^\frac{1}{2} $$



with n = the number of observations

Calculate the rmse and save the result as `rmse` (not in your Dataframe):

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In [20], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
rmse
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [21], line 1
----> 1 rmse

NameError: name 'rmse' is not defined
```

<!-- grade_id: c-rmse, points: 2, grade: true, solution: false -->
```python
# Check your code
assert 1203 < rmse < 1206
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [22], line 2
      1 # Check your code
----> 2 assert 1203 < rmse < 1206

NameError: name 'rmse' is not defined
```

### Mean absolute error

Another performance measure (which is less commonly used) is the mean absolute error:

The general formula:

$$Mean\ absolute\ error \ (MAE) = \frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})$$



### Use functions

#### Mean squared error

```python
mean_squared_error(df['sales'], df['sales_prediction'])
```

Output:
```
1452000.0
```

#### Root mean squared error

```python
mean_squared_error(df['sales'], df['sales_prediction'], squared=False)
```

Output:
```
1204.9896265113655
```

#### Mean absolute error

```python
mean_absolute_error(df['sales'], df['sales_prediction'])
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [25], line 1
----> 1 mean_absolute_error(df['sales'], df['sales_prediction'])

NameError: name 'mean_absolute_error' is not defined
```