**Important notes:**

**Important notes:**

- You can complete this exercise in 1) Colab or 2) on your local machine
  1. **Colab**: Click on the 🚀 symbol at the top of the page and select Colab. When you finished the exercise, download the file: `File` > `Download` > `ipynb`. 
  2. **Local**: click on the download button at the top of the page and choose `.ipynb`. Activate the conda environment `mr` before you start: `conda activate mr`


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.


- The function **NotImplementedError()** prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - in *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - if you use *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 


Good luck!

```python
NAME = "ar166"
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Find the best model

## Setup

```python
import pandas as pd

from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
from sklearn.metrics import mean_absolute_error
```

## Data

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

### Variable lists

We need to prepare our data for the scikit-learn model (which does not work with Pandas dataframes):

- We create the variable `y_label` (our dependent variable, also called label): choose "`sales`" (this step will be helpful in later notebooks)

- We select our features (independent variables) and save them as `X`: choose "`ads`"

- Select and save the dependent variable as `y`

Hint:

```pyhton
y_label = "___"

X = df[["___"]]
y = df[___]
```

<!-- grade_id: y_label, points: 0, grade: false, solution: true -->
```python
y_label  = 'sales'


```

```python
X = df[['ads']]


```

```python
y = df['y_label']
```

Output:
```
File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/frame.py:3805, in DataFrame.__getitem__(self, key)
   3803 if self.columns.nlevels > 1:
   3804     return self._getitem_multilevel(key)
-> 3805 indexer = self.columns.get_loc(key)
   3806 if is_integer(indexer):
   3807     indexer = [indexer]

File ~/opt/anaconda3/envs/nb/lib/python3.9/site-packages/pandas/core/indexes/base.py:3802, in Index.get_loc(self, key, method, tolerance)
   3800     return self._engine.get_loc(casted_key)
   3801 except KeyError as err:
-> 3802     raise KeyError(key) from err
   3803 except TypeError:
   3804     # If we have a listlike key, _check_indexing_error will raise
   3805     #  InvalidIndexError. Otherwise we fall through and re-raise
   3806     #  the TypeError.
   3807     self._check_indexing_error(key)

KeyError: 'y_label'
```

```python
X
```

Output:
```
ads
0    900
1   1400
2   3600
3   3800
4   6200
5   5200
6   6800
7   8300
8   9800
9  10100
```

```python
y
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [9], line 1
----> 1 y

NameError: name 'y' is not defined
```

<!-- grade_id: c, points: 2, grade: true, solution: false -->
```python
### Check your code (note that there are some hidden checks)
assert len(X) == 10

### BEGIN HIDDEN TESTS 
assert y_label == "sales"
assert len(y) == 10
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [10], line 6
      4 ### BEGIN HIDDEN TESTS 
      5 assert y_label == "sales"
----> 6 assert len(y) == 10

NameError: name 'y' is not defined
```

## Model

### Select model

Hint:

Save the `LinearRegression()`estimator as object called `reg`

---

```pyhton
___ = ___()
```

---

<!-- grade_id: reg, points: 0, grade: false, solution: true -->
```python
reg = LinearRegression()
```

<!-- grade_id: c-reg, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.get_params() == {'copy_X': True, 'fit_intercept': True, 'n_jobs': None, 'normalize': 'deprecated', 'positive': False}
```

### Fit model


Fit the model (your `reg` object) to the data with `.fit()`.

In the fit function, you first need to provide the features and then the dependent variable (features, dependent variable)

Hint:


---

```pyhton
___.___(___, ___)
```

---



<!-- grade_id: reg-fit, points: 0, grade: false, solution: true -->
```python
reg.fit(X, y)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [13], line 1
----> 1 reg.fit(X, y)

NameError: name 'y' is not defined
```

<!-- grade_id: c-reg-fit, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.positive == False
```

### Coefficients

```python
# Intercept
reg.intercept_
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In [15], line 2
      1 # Intercept
----> 2 reg.intercept_

AttributeError: 'LinearRegression' object has no attribute 'intercept_'
```

```python
# Slope
reg.coef_
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In [16], line 2
      1 # Slope
----> 2 reg.coef_

AttributeError: 'LinearRegression' object has no attribute 'coef_'
```

### Make predictions


Use your model to:

- predict sales from your ads (`X`) values with the method `.predict()`

- Save the result as `y_pred`

Hint:

---

```pyhton
___ = ___.___(___)
```

---



<!-- grade_id: y-pred, points: 0, grade: false, solution: true -->
```python
y_pred = x.predict()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [17], line 1
----> 1 y_pred = x.predict()

NameError: name 'x' is not defined
```

```python
y_pred[0]
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [18], line 1
----> 1 y_pred[0]

NameError: name 'y_pred' is not defined
```

<!-- grade_id: c-y-pred, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 2836 < y_pred[0] <  2840
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [19], line 2
      1 ### Check your code
----> 2 assert 2836 < y_pred[0] <  2840

NameError: name 'y_pred' is not defined
```

### Evaluation

#### Mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the mean squared error and save the result as `mse`

Hint:

---

```pyhton
___ = ___(___, ___)
```

---

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
mse = mean_squared_error(x, y)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [20], line 1
----> 1 mse = mean_squared_error(x, y)

NameError: name 'x' is not defined
```

```python
print(f'The mean squared error equals {mse:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [21], line 1
----> 1 print(f'The mean squared error equals {mse:.0f} ')

NameError: name 'mse' is not defined
```

<!-- grade_id: c-mse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1160738 < mse < 1160740 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [22], line 2
      1 ### Check your code
----> 2 assert 1160738 < mse < 1160740

NameError: name 'mse' is not defined
```

#### Root mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the root mean squared error and save the result as `rmse`.

Hint:

---

```pyhton
___ = ___(___, ___, ___)
```

- use the option `squared=False` 

---

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
rmse = mean_squared_error(x, y, mse)
squared=False
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [23], line 1
----> 1 rmse = mean_squared_error(x, y, mse)
      2 squared=False

NameError: name 'x' is not defined
```

```python
print(f'The root mean squared error equals {rmse:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [24], line 1
----> 1 print(f'The root mean squared error equals {rmse:.0f} ')

NameError: name 'rmse' is not defined
```

<!-- grade_id: c-rmse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1076 < rmse < 1078 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [25], line 2
      1 ### Check your code
----> 2 assert 1076 < rmse < 1078

NameError: name 'rmse' is not defined
```

#### Mean absolute error

Use the scikit-learn function `mean_absolute_error()` to calculate the mean absolute error and save the result as `mae`.

Hint:

---

```pyhton
___ = ___(___, ___)
```


---

<!-- grade_id: mae, points: 0, grade: false, solution: true -->
```python
mae = mean_absolute_error(x, y)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [26], line 1
----> 1 mae = mean_absolute_error(x, y)

NameError: name 'x' is not defined
```

```python
print(f'The mean absolute error equals {mae:.0f} ')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [27], line 1
----> 1 print(f'The mean absolute error equals {mae:.0f} ')

NameError: name 'mae' is not defined
```

<!-- grade_id: c-mae, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 885 < mae < 887 
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In [28], line 2
      1 ### Check your code
----> 2 assert 885 < mae < 887

NameError: name 'mae' is not defined
```